---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction to RNN and Basic Assumptions of various models

```yaml
type: VideoExercise
key: 60a73b9ef2
xp: 50
```

`@projector_key`
338dbda74c890e5c14f63e9305412f7d
